rule_files:
  - /etc/prometheus/alerts.yml
evaluation_interval: 1m
tests:
- interval: 1m
  input_series:
  - series: 'chs_predict_errors_total{model_id="default",model_version="0"}'
    values: '0 1 2 3 4 5 6 7 8 9 10 11'
  - series: 'chs_predict_latency_ms_bucket{le="150",model_id="default",model_version="0"}'
    values: '0 0 0 0 0 0 0 0 0 0 0 0'
  - series: 'chs_predict_latency_ms_bucket{le="250",model_id="default",model_version="0"}'
    values: '0 10 20 30 40 50 60 70 80 90 100 110'
  - series: 'chs_predict_latency_ms_bucket{le="+Inf",model_id="default",model_version="0"}'
    values: '0 10 20 30 40 50 60 70 80 90 100 110'
  - series: 'chs_model_reload_failures_total{model_id="default",model_version="0",reason="error"}'
    values: '0 1 2 3 4 5 5 5 5 5 5 5'
  alert_rule_test:
  - eval_time: 6m
    alertname: PredictHighErrorRate
    exp_alerts:
    - exp_labels:
        severity: warning
        team: obs
        model_id: default
        model_version: "0"
      exp_annotations:
        summary: High predict error rate
        description: Error rate >20% for 5m
  - eval_time: 11m
    alertname: PredictLatencyP95High
    exp_alerts:
    - exp_labels:
        severity: warning
        team: obs
        model_id: default
        model_version: "0"
      exp_annotations:
        summary: Predict latency P95 high
        description: P95 latency above 150ms for 10m
  - eval_time: 6m
    alertname: ModelReloadFailuresSpike
    exp_alerts:
    - exp_labels:
        severity: critical
        team: obs
        model_id: default
        model_version: "0"
        reason: error
      exp_annotations:
        summary: Model reload failures spike
        description: More than 3 reload failures within 30m
